# Algorithms Exercise 2

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import seaborn as sns
import numpy as np
```

## Peak finding

Write a function `find_peaks` that finds and returns the indices of the local maxima in a sequence. Your function should:

* Properly handle local maxima at the endpoints of the input array.
* Return a Numpy array of integer indices.
* Handle any Python iterable as input.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def find_peaks(a):
    """Find the indices of the local maxima in a sequence."""
    maxima = []
    length = len(a)
    if length >= 2:
        if a[0] > a[1]:
            p = 0
            maxima.append(p)
            
    if length > 3:
        for i in a:
            if a[i] > a[i-1] and a[i] > a[i+1]:
                q = np.nonzero(a[i])
                maxima.append(q)
    
    if a[length-1] > a[length-2]:
        r = a[length-1]
        maxima.append(r)        
    return maxima


find_peaks([2,0,1,0,2,0,1])

        
    
```

Output:
```
[0,
 (array([0]),),
 (array([0]),),
 (array([0]),),
 (array([0]),),
 (array([0]),),
 1]
```

```python
a = [1,3,6,4,5,1]
find_peaks(a)
```

Output:
```
<ipython-input-109-494ba1fc916d> in <module>()
      1 a = [1,3,6,4,5,1]
----> 2 find_peaks(a)

<ipython-input-108-1d08c3d0025f> in find_peaks(a)
     10     if length > 3:
     11         for i in a:
---> 12             if a[i] > a[i-1] and a[i] > a[i+1]:
     13                 q = np.argmax(a)
     14                 maxima.append(q)

IndexError: list index out of range
```

<!-- grade_id: algorithmsex02a, points: 5, grade: true, solution: false -->
```python
p1 = find_peaks([2,0,1,0,2,0,1])
assert np.allclose(p1, np.array([0,2,4,6]))
p2 = find_peaks(np.array([0,1,2,3]))
assert np.allclose(p2, np.array([3]))
p3 = find_peaks([3,2,1,0])
assert np.allclose(p3, np.array([0]))
```

Output:
```
<ipython-input-87-e9503c7dbdc3> in <module>()
      1 p1 = find_peaks([2,0,1,0,2,0,1])
----> 2 assert np.allclose(p1, np.array([0,2,4,6]))
      3 p2 = find_peaks(np.array([0,1,2,3]))
      4 assert np.allclose(p2, np.array([3]))
      5 p3 = find_peaks([3,2,1,0])

/usr/local/lib/python3.4/dist-packages/numpy/core/numeric.py in allclose(a, b, rtol, atol)
   2232     # ignore invalid fpe's
   2233     with errstate(invalid='ignore'):
-> 2234         r = all(less_equal(abs(x - y), atol + rtol * abs(y)))
   2235 
   2236     return r

ValueError: operands could not be broadcast together with shapes (7,) (4,)
```

Here is a string with the first 10000 digits of $\pi$ (after the decimal). Write code to perform the following:

* Convert that string to a Numpy array of integers.
* Find the indices of the local maxima in the digits of $\pi$.
* Use `np.diff` to find the distances between consequtive local maxima.
* Visualize that distribution using an appropriately customized histogram.

```python
from sympy import pi, N
pi_digits_str = str(N(pi, 10001))[2:]
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: algorithmsex02b, points: 5, grade: true, solution: false -->
```python
assert True # use this for grading the pi digits histogram
```