# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    q = list(filter(lambda x: x in y, y))
    p = list(filter(lambda y: y in x, x)) 

    
    for i in q:
        if i in p and len(q) == len(p):
            y = True
        else:
            y = False
    return y
```

```python
#x ='132'; y = '321'
#
#q = list(filter(lambda x: x in y, y))
#p = list(filter(lambda y: y in x, x)) 
#
#for i in q:
#    if i in p:
#        y = True
#    else:
#        y = False
#print(y)

#
#print(q)
#print(p)
```

Output:
```
File "<ipython-input-72-9589dcc34a89>", line 9
    else if:
          ^
SyntaxError: invalid syntax
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
xx=range(1,100)
yy=range(1,100)

pal_list =[]
for xnum in xx:
    for ynum in yy:
        pal_list.append((xnum, ynum))
        
list0 = str(pal_list)

# I found the list of the (x,y) combination for range 1 - 100.... I ran out of time :(
```

Output:
```
File "<ipython-input-99-3c6e3bf9730d>", line 10
    list0 = as.array(list0)
             ^
SyntaxError: invalid syntax
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```